---
title       : Review of Statistics
description : Insert the chapter description here
---
## Biased estimator

```yaml
type: NormalExercise
key: 76c959aa9b
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---
## <<<New Exercise>>>

```yaml
type: NormalExercise
key: b75a7ad1f3
lang: r
xp: 100
skills: 1
```


`@instructions`

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}

```

`@sct`
```{r}

```

---
## Efficiency of an estimator

```yaml
type: NormalExercise
key: ac77115a16
lang: r
xp: 100
skills: 1
```

In this exercise we want to illustrate that the 

`@instructions`

- Verify that the alternative estimator $\tilde{Y}$

`@hint`

`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}
# Verify that the alternative estimator is unbiased as well
w <- c(rep((1+0.5)/100, 50), rep((1-0.5)/100, 50))

```

`@solution`
```{r}

```

`@sct`
```{r}

```
